## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc,  to develop a new version of the function with a simpler interface (the interface is the arguments that are passed to the function). Add appropriate checks inside your function to validate the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def triangleArea(a,b,c):
    "return area of triangle in 2D"
    
    #Validate Data (ensure each vector has 2 elements only)
    if len(a) == 2 and len(b) == 2 and len(c) == 2:
        if (type(a) == float or type(a) == int or type(a) == complex) and type(a) == type(b) and type(a) == type(c):
            pass
    else:
        print("Error: each vector must have 2 components only")
        return None
    
    #Prev code: return abs(x0*(y1-y2) + x1*(y2-y0) + x2*(y0 - y1))/2
    #Replace all instance of x,y,0,1,2 with [0],[1],a,b,c
    return abs(a[0]*(b[1]-c[1]) + b[0]*(c[1]-a[1]) + c[0]*(a[1] - b[1]))/2
```

```python
#Test area of triangle form Ex 4

a = [0.0 + 0j, 0.0]
b = [0.0, 2.0]
c = [3.0, 0.0]
A = triangleArea(a,b,c)
print(A) # Should = 3.0

```

Output:
```
3.0
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that computes the area of a simple polygon with an arbitrary number of vertices. Test your function for some simple shapes. Pay careful attention to the range of any loops.

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def simple_polygon_area(vertices):
    n = len(vertices)
    vertices.append(vertices[0]) # Add first element to last index, so that vertices[n] exists
    tot = 0
    for i in range (n):
        tot += vertices[i][0] * vertices[i+1][1] - vertices[i+1][0] * vertices[i][1]
    tot /= 2
    return abs(tot)
```

```python
#Test for triangle
a = (1.0, 1.0)
b = (3.0, 1.0)
c = (3.0, 3.0)
d = (1.0, 3.0)
ver = [a,b,c,d]
print(simple_polygon_area(ver))

# #Test for square
# d = (3., 2.)
# ver = [a,b,d,c]
# print(simple_polygon_area(ver))

```

Output:
```
4.0
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match, and print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n
    
#### Optional (advanced) 

Try writing a one-line version of this operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y):
    "Return sum of two vectors"
    # YOUR CODE HERE
    n = len(x)
    if n == len(y): # Check if x has the same length as y
        tot = [0]*n
        for i in range (n):
            tot[i] = x[i] + y[i]
    else:
        print("Error: x does not have the same length as y")
        return None
    
    if n == len(tot):
        return tot
    else:
        print("Error: tot does not have the same length as x")
        return None
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Extension: list comprehension

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
d = [a[i] + b[i] for i in range (len(a))]
print(d)
```

Output:
```
[-2, 11.3, -20, 8]
```

```python
[sum(c) for c in zip(a,b)]
```

Output:
```
[-2, 11.3, -20, 8]
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A dictionary from college abbreviation to name; and
1. A list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation 
- Total number students
 
for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary, 

1. Find the college with the greatest number of students and print the abbreviation; and 
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
college_name_to_abb = {"Christ's": ["CHR"], "Churchill": ["CHU"], "Clare": ["CL"], "Clare Hall": ["CLH"], "Corpus Christi": ["CC"], "Dawrin": ["DAR"], "Downing": ["DOW"], "Emmanuel": ["EM"], 
                       "Fitzwilliam": ["F"], "Girton": ["G"], "Gonville and Caius": ["CAI"], "Homerton": ["HO"], "Hughes Hall": ["HH"], "Jesus": ["JE"], "King's": ["K"], "Lucy Cavendish": ["LC"], 
                       "Magdalene": ["M"], "Murray Edwards": ["MUR"], "Newnham": ["N"], "Pembroke": ["PEM"], "Peterhouse": ["PET"], "Queens'": ["Q"], "Robinson": ["R"], "St Catherine's": ["CTH"], 
                       "St Edmund's": ["ED"], "Selwyn": ["SE"], "Sidney Sussex": ["SID"], "Trinity": ["T"], "Trinity Hall": ["TH"], "Wolfson": ["W"],}

# 1. Create a dictionary from college abbreviation to name:
college_abb_to_name = {}
for key, value in college_name_to_abb.items():
    college_abb_to_name[value[0]] = key

print(college_abb_to_name)

# 2. A list of college abbreviations sorted in alphabetical order
college_abb = []
for abb in college_abb_to_name:
    college_abb.append(abb)

college_abb.sort()
print(college_abb)

```

Output:
```
{'DOW': 'Downing', 'DAR': 'Dawrin', 'CTH': "St Catherine's", 'CAI': 'Gonville and Caius', 'K': "King's", 'SE': 'Selwyn', 'JE': 'Jesus', 'W': 'Wolfson', 'MUR': 'Murray Edwards', 'G': 'Girton', 'ED': "St Edmund's", 'HO': 'Homerton', 'LC': 'Lucy Cavendish', 'Q': "Queens'", 'T': 'Trinity', 'CHR': "Christ's", 'CL': 'Clare', 'F': 'Fitzwilliam', 'N': 'Newnham', 'HH': 'Hughes Hall', 'CLH': 'Clare Hall', 'SID': 'Sidney Sussex', 'M': 'Magdalene', 'CC': 'Corpus Christi', 'CHU': 'Churchill', 'TH': 'Trinity Hall', 'R': 'Robinson', 'EM': 'Emmanuel', 'PEM': 'Pembroke', 'PET': 'Peterhouse'}
['CAI', 'CC', 'CHR', 'CHU', 'CL', 'CLH', 'CTH', 'DAR', 'DOW', 'ED', 'EM', 'F', 'G', 'HH', 'HO', 'JE', 'K', 'LC', 'M', 'MUR', 'N', 'PEM', 'PET', 'Q', 'R', 'SE', 'SID', 'T', 'TH', 'W']
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
college_name_to_abb = {"Christ's": ["CHR",1505,698], "Churchill": ["CHU",1960,704], "Clare": ["CL",1326,655], "Clare Hall": ["CLH",1966,155], "Corpus Christi": ["CC",1352,467], 
                       "Dawrin": ["DAR",1964,710], "Downing": ["DOW",1800,771], "Emmanuel": ["EM",1584,750], "Fitzwilliam": ["F",1869,855], "Girton": ["G",1869,776], 
                       "Gonville and Caius": ["CAI",1348,786], "Homerton": ["HO",1768,1118], "Hughes Hall": ["HH",1885,736], "Jesus": ["JE",1496,911], "King's": ["K",1441,648], 
                       "Lucy Cavendish": ["LC",1965,380], "Magdalene": ["M",1428,529], "Murray Edwards": ["MUR",1954,580], "Newnham": ["N",1871,638], "Pembroke": ["PEM",1347,597], 
                       "Peterhouse": ["PET",1284,370], "Queens'": ["Q",1448,1034], "Robinson": ["R",1977,495], "St Catherine's": ["CTH",1473,621], "St Edmund's": ["ED",1896,574], 
                       "Selwyn": ["SE",1882,626], "Sidney Sussex": ["SID",1596,602], "Trinity": ["T",1546,1061], "Trinity Hall": ["TH",1350,612], "Wolfson": ["W",1965,887],}

# 1. Find the college with the greatest number of students and print the abbreviation
greatest_num_students = 0
college_greatest = ""
for key, value in college_name_to_abb.items():
    if value[2] > greatest_num_students:
        greatest_num_students = value[2]
        college_greatest = key
print("College with greatest number of students is {} with {} students".format(college_name_to_abb[college_greatest][0], greatest_num_students))

# 2. Find the oldest college, and print the number of students and the abbreviation for this college
oldest = 2019
college_oldest = ""
for key, value in college_name_to_abb.items():
    if value[1] < oldest:
        oldest = value[1]
        college_oldest = key
print("Oldest College is {} with {} students. Age is {}".format(college_name_to_abb[college_oldest][0], greatest_num_students, college_name_to_abb[college_oldest][1]))

```

Output:
```
College with greatest number of students is HO with 1118 students
Oldest College is PET with 1118 students. Age is 1284
```

```python

```