## Exercise 09.1 (checking data validity)

The Fibonacci series is valid only for $n \ge 0$. Add to the Fibonacci function in this notebook a check that raises an exception if $n < 0$. Try some invalid data cases to check that an exception is raised.

*Optional:* Use `pytest` to test that an exception *is* raised for some $n < 0$ cases.

<!-- grade_id: cell-982c3611466532a2, points: 0, grade: false, solution: true -->
```python
def f(n): 
    "Compute the nth Fibonacci number using recursion"
    if n == 0:
        return 0
    elif n == 1:
        return 1
    elif n < 0:
        raise ValueError("Entered number must be positive")
    else:
        return f(n-1)+f(n-2)
```

<!-- grade_id: cell-9b97ea3e85d77f4a, points: 0, grade: true, solution: false -->
```python
# Perform some tests    
assert f(0) == 0
assert f(1) == 1
assert f(2) == 1
assert f(3) == 2
assert f(10) == 55
assert f(15) == 610

# Check that ValueError is raised for n < 0
import pytest
with pytest.raises(ValueError):
    f(-1)
with pytest.raises(ValueError):
    f(-2)
```

## Exercise 09.2 (raising exceptions)

Modify your program from the bisection exercise in Activity 04 to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

<!-- grade_id: cell-ca2c3ee62ad3b956, points: 0, grade: false, solution: true -->
```python
def my_f(x):
    "Evaluate polynomial function"
    return x**3 - 6*x**2 + 4*x + 12

def compute_root(f, x0, x1, tol, max_it):
    "Compute roots of a function using bisection"
    #Perform basic error handling
    #Check if there is a sign change between x0 and x1 to ensure that the root lies between these values
    if f(x0)*f(x1) > 0: #Check if the result is positive, hence there is no sign change
        raise ValueError("No sign change between x0 and x1, hence the root doesn't lie inbetween these values or there's a TP nearby")
    
    it = 0
    error = tol + 1. #Set error to be greater than the tolerance
    
    while error > tol and it < max_it:
        #Increment counter
        it += 1
        
        #Get mid val
        x_mid = (x0 + x1)/2
        
        #Compute f(x0) and f(x_mid)
        f0 = f(x0)
        f_mid = f(x_mid)
        
        #Get the sign
        sign = f0 * f_mid
        
        if sign < 0: #Negative, so move x1 to x_mid
            x1 = x_mid
        else:
            x0 = x_mid
        
        #Get error
        error = abs(f_mid)
    
    if it >= max_it: #Raise  error if the iteration has exceeded the limit
        raise RecursionError("Number of iterations has exceeded {}".format(max_it)) #Is RuntimeError better?
    
    return x_mid, f_mid, it
```

<!-- grade_id: cell-f6cf91ed5dd1736b, points: 0, grade: true, solution: false -->
```python
# Test with max_it = 30
x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=30)

# Test with max_it = 20
with pytest.raises(RuntimeError):
    x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=20)
```

```python
with pytest.raises(ValueError):
    x, f, num_it = compute_root(my_f, x0=5, x1=6, tol=1.0e-6, max_it=20)
```